Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# CSE 30 Fall 2021 - Homework 4


### Instructions

Please disregard the YOUR NAME and COLLABORATORS above.  They are put there atomatically by the grading tool.
You can find instructions on how to work on a homework on Canvas.  Here is a short summary: 

### Submitting your work

To submit your work: 

* First, click on "Runtime > Restart and run all", and check that you get no errors.  This enables you to catch any error you might have introduced, and not noticed, due to your running cells out of order. 
* Second, download the notebook in .ipynb format (File > Download .ipynb) and upload the .ipynb file to [this form](https://docs.google.com/forms/d/e/1FAIpQLSe7oq8KxzdbDKmeKBWaK--qBF04zuXcqumX8B3KikU_RncWPg/viewform?usp=sf_link). This homework is due by 11:59pm on Friday 21 January 2022.

You can submit multiple times; the last submission before the deadline is the one that counts.

### Homework format

For each question in this notebook, there is: 

* A text description of the problem. 
* One or more places where you have to insert your solution.  You need to complete every place marked: 

    `# YOUR CODE HERE`
    
    and you should not modify any other place. 
* One or more test cells.  Each cell is worth some number of points, marked at the top.  You should not modify these tests cells.  The tests pass if no error is printed out: when there is a statement that says, for instance: 

    `assert x == 2`
    
    then the test passes if `x` has value 2, and fails otherwise.  You can insert a `print(x)` (for this case!) somewhere if you want to debug your work; it is up to you.  
    
### Notes:

* Your code will be tested both according to the tests you can see (the `assert` statements you can see), _and_ additional tests.  This prevents you from hard-coding the answer to the particular questions posed.  Your code should solve the _general_ intended case, not hard-code the particular answer for the values used in the tests. 

* **Please do not delete or add cells!** The test is autograded, and if you modify the test by adding or deleting cells, even if you re-add cells you delete, you may not receive credit. 

* **Please do not import modules that are not part of the [standard library](https://docs.python.org/3/library/index.html).** You do not need any, and they will likely not available in the grading environment, leading your code to fail. 

* **If you are inactive too long, your notebook might get disconnected from the back-end.** Your work is never lost, but you have to re-run all the cells before you continue. 

* You can write out print statements in your code, to help you test/debug it. But remember: the code is graded on the basis of what it outputs or returns, not on the basis of what it prints.

* **TAs and tutors have access to this notebook,** so if you let them know you need their help, they can look at your work and give you advice. 

### Grading

Each cell where there are tests is worth a certain number of points.  You get the points allocated to a cell only if you pass _all_ the tests in the cell. 

The tests in a cell include both the tests you can see, and other, similar, tests that are used for grading only.  Therefore, you cannot hard-code the solutions: you really have to solve the essence of the problem, to receive the points in a cell. 

### Code of Conduct

* Work on the test yourself, alone. 
* You can search documentation on the web, on sites such as the Python documentation sites, Stackoverflow, and similar, and you can use the results. 
* You cannot share your work with others or solicit their help.


## Problem 1 
### Write a generator for the Fibonacci Numbers

Build a generator that returns the [Fibonacci numbers](https://en.wikipedia.org/wiki/Fibonacci_number): 0, 1, 1, 2, 3, 5, and so on. 

<!-- grade_id: cell-272a42f4ce1df5f9, points: 0, grade: false, solution: true -->
```python
def fibonacci_generator():
    """Generates all Fibonacci numbers."""
    # YOUR CODE HERE

    a, b = 0, 1
    while True:
        yield a
        a, b = b, a + b
'''
    list = []

    for j in range(1, 100):
        result = 0
        f1, f2 = 0, 1
        for i in range(0, j):
            if i <= 1:
                result = i
            else:
                result = f1 + f2
            f1 = f2
            f2 = result
        list.append(result)

    return list

    def get():
    n = 100_000
    mylist = [0, 1]
    [mylist.append(mylist[-2] + mylist[-1]) for _ in range(n)]
    return mylist
'''
        
```

Output:
```
'\n    list = []\n\n    for j in range(1, 100):\n        result = 0\n        f1, f2 = 0, 1\n        for i in range(0, j):\n            if i <= 1:\n                result = i\n            else:\n                result = f1 + f2\n            f1 = f2\n            f2 = result\n        list.append(result)\n\n    return list\n\n    def get():\n    n = 100_000\n    mylist = [0, 1]\n    [mylist.append(mylist[-2] + mylist[-1]) for _ in range(n)]\n    return mylist\n'
```

```python
## Here you can test your code.
n = 0
fib = fibonacci_generator()
while n < 100:
    print(next(fib))
    n += 1

# print(', '.join(str(next(fib)) for _ in range()))
## What you write here will be removed.

```

Output:
```
0
1
1
2
3
5
8
13
21
34
55
89
144
233
377
610
987
1597
2584
4181
6765
10946
17711
28657
46368
75025
121393
196418
317811
514229
832040
1346269
2178309
3524578
5702887
9227465
14930352
24157817
39088169
63245986
102334155
165580141
267914296
433494437
701408733
1134903170
1836311903
2971215073
4807526976
7778742049
12586269025
20365011074
32951280099
53316291173
86267571272
139583862445
225851433717
365435296162
591286729879
956722026041
1548008755920
2504730781961
4052739537881
6557470319842
10610209857723
17167680177565
27777890035288
44945570212853
72723460248141
117669030460994
190392490709135
308061521170129
498454011879264
806515533049393
1304969544928657
2111485077978050
3416454622906707
5527939700884757
8944394323791464
14472334024676221
23416728348467685
37889062373143906
61305790721611591
99194853094755497
160500643816367088
259695496911122585
420196140727489673
679891637638612258
1100087778366101931
1779979416004714189
2880067194370816120
4660046610375530309
754011380474634642
```

```python
# Visible tests: 5 points

r = []
for n in fibonacci_generator():
    r.append(n)
    if n > 100:
        break
assert r == [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144]

```

<!-- grade_id: cell-1e461d53e390cba5, points: 10, grade: true, solution: false -->
```python
# Hidden tests, 5 points

# These tests check basically that you generate Fibonacci numbers, forever -- without stopping.
# It is the reader (the user) of the function that decides when to stop.

c = 0
for n in fibonacci_generator():
    if n > 100:
        c += 1
    if n in [32, 88, 99, 231, 1000]:
        print("bad:", n)
        raise Exception()
    if n > 1000:
        break
assert c == 6


```

## Problem 2
### Write a prime number generator

Write a generator that returns all the prime numbers.  The idea is to loop over all positive integers, test each one to see if it is prime, and if it is, `yield` it. 

<!-- grade_id: cell-6486f8e9d1595cc6, points: 0, grade: false, solution: true -->
```python
# My solution is simple and not particularly optimized,
# and it is 12 lines long.
import math 
def prime_number_generator():
    """This generator returns all prime numbers."""
    # YOUR CODE HERE

    dict = {}
    x = 2
    while True: 
        if x not in dict:
            yield x
            dict[x*x] = [x]
        else:
            for i in dict[x]:
                dict.setdefault(i+x, []).append(i)
            del dict[x]
        x += 1

```

```python
## Here you can test your code.
initial_primes = [2, 3, 5, 7, 11, 13, 17, 19, 23]
r = []

for n in prime_number_generator():
    r.append(n)
    if n >= 300:
        break
print(r)
## What you write here will be removed.

```

Output:
```
[2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71, 73, 79, 83, 89, 97, 101, 103, 107, 109, 113, 127, 131, 137, 139, 149, 151, 157, 163, 167, 173, 179, 181, 191, 193, 197, 199, 211, 223, 227, 229, 233, 239, 241, 251, 257, 263, 269, 271, 277, 281, 283, 293, 307]
```

```python
## Visible tests: 5 points

initial_primes = [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71, 73, 79, 83, 89, 97]

idx = 0
for p in prime_number_generator():
    assert p == initial_primes[idx]
    idx += 1
    if idx == len(initial_primes):
        break
    print(p)
```

Output:
```
2
3
5
7
11
13
17
19
23
29
31
37
41
43
47
53
59
61
67
71
73
79
83
89
```

<!-- grade_id: cell-a7e50a3c136cec77, points: 10, grade: true, solution: false -->
```python
## Hidden tests: 10 points

# We test that the generator goes on and on and on.
c = 0
for n in prime_number_generator():
    if n > 900:
        c += 1
    if n in [32, 88, 99, 231, 1000]:
        print("bad:", n)
        raise Exception()
    if n > 1000:
        break
assert c == 15

```